# Métodos Matemáticos Computacionales para Ciencia de Datos (CdD 0503)
# Taller 01 — Métodos numéricos para sistemas lineales y factorizaciones
### Facultad de Estudios Superiores Acatlán · Universidad Nacional Autónoma de México

**FES Acatlán · UNAM**  
**Duración sugerida:** 60–75 minutos  
**Puntaje total:** 40 puntos

## Instrucciones

- No agregue ni elimine celdas.
- No modifique las celdas de pruebas.
- Escriba código únicamente en las celdas marcadas con `# ESCRIBA SU CÓDIGO AQUÍ`.
- Ejecute las celdas en orden.
- Antes de entregar, reinicie el kernel y ejecute todo el notebook.
- Las pruebas visibles verifican aspectos básicos; también existen pruebas ocultas de corrección numérica.

## Temas

- Jacobi y Gauss-Seidel.
- Factorización LU por Crout.
- Factorización de Cholesky.
- Factorización QR mediante Gram-Schmidt modificado.
---

### Nombre del estudiante

**Nombre:** __________________________________________

**Número de cuenta:** _________________________________

**Fecha:** ____________________________________________

```python
import numpy as np
np.set_printoptions(precision=6, suppress=True)
```

## Pregunta 1 — Una iteración de Jacobi y Gauss-Seidel (10 puntos)

Implemente las funciones `jacobi_step(A, b, x)` y `gauss_seidel_step(A, b, x)` que realicen **exactamente una iteración** de cada método.

Para Jacobi, todas las componentes nuevas deben calcularse usando únicamente el vector anterior $x^{(k)}$:

$$
x_i^{(k+1)}
=
\frac{1}{a_{ii}}
\left(
b_i-\sum_{j\neq i}a_{ij}x_j^{(k)}
\right).
$$

Para Gauss-Seidel, use inmediatamente las componentes nuevas ya calculadas:

$$
x_i^{(k+1)}
=
\frac{1}{a_{ii}}
\left[
b_i
-\sum_{j=1}^{i-1}a_{ij}x_j^{(k+1)}
-\sum_{j=i+1}^{n}a_{ij}x_j^{(k)}
\right].
$$

No use `np.linalg.solve`.

<!-- grade_id: q1_solution, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: q1_tests, points: 10, grade: true, solution: false -->
```python
A1 = np.array([[4., -1., 0.],
               [-1., 4., -1.],
               [0., -1., 3.]])
b1 = np.array([15., 10., 10.])
x0 = np.zeros(3)

j1 = jacobi_step(A1, b1, x0)
g1 = gauss_seidel_step(A1, b1, x0)

assert isinstance(j1, np.ndarray)
assert isinstance(g1, np.ndarray)
assert j1.shape == (3,)
assert g1.shape == (3,)


```

## Pregunta 2 — Factorización LU por Crout (10 puntos)

Implemente `crout(A)` para calcular

$
A=LU,
$

donde $L$ es triangular inferior y $U$ es triangular superior con diagonal unitaria:

$
u_{ii}=1.
$

Suponga que la matriz admite la factorización sin pivoteo y que no aparecen pivotes nulos.

La función debe regresar `L, U`.

## Código Guía
```Python
def crout(A):
    A = np.asarray(____________, dtype=float)
    n = A.shape[0]
    L = np.zeros_like(____________, ____________)
    U = np.eye(n, dtype=float)

    for j in range(____________):
        for i in range(____________, ____________):
            L[i, j] = A[i, j] - np.dot(L[i, :j], U[:j, j])

        for i in range(j + 1, n):
            U[j, i] = (A[j, i] - np.dot(L[j, :j], U[:j, i])) / L[j, j]

    return ____________, ____________
```

<!-- grade_id: q2_solution, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: q2_tests, points: 10, grade: true, solution: false -->
```python
A2 = np.array([[2., 2., 0.],
               [4., 7., 6.],
               [-2., -1., 7.]])

L2, U2 = crout(A2)

assert L2.shape == A2.shape
assert U2.shape == A2.shape
assert np.allclose(np.triu(L2, 1), 0)
assert np.allclose(np.tril(U2, -1), 0)
assert np.allclose(np.diag(U2), 1)


```

## Pregunta 3 — Cholesky y solución de un sistema (10 puntos)

Implemente `cholesky_solve(A, b)` para una matriz simétrica definida positiva.

Debe:

1. construir $L$ tal que
$
A=LL^T,
$
2. resolver
$
Ly=b,
$
3. resolver
$
L^Tx=y,
$
4. regresar `L, x`.

**Restricción:** no use `np.linalg.cholesky`, `np.linalg.solve` ni `np.linalg.inv`.

## Código Guía
```Python 
def cholesky_solve(A, b):
    A = np.asarray(________, ________)
    b = np.asarray(________, ________)
    n = A.shape[0]
    L = np.zeros_like(A, dtype=float)

    # Factorización de Cholesky
    for i in range(n):
        for j in range(________):
            s = np.dot(L[i, :j], L[j, :j])
            if i == j:
                L[i, j] = np.sqrt(A[i, i] - s)
            else:
                L[i, j] = (A[i, j] - s) / L[j, j]

    # Sustitución hacia adelante: Ly = b
    y = np.zeros(n, dtype=float)
    for i in range(n):
        y[i] = (b[i] - np.dot(L[i, :i], y[:i])) / L[i, i]

    # Sustitución hacia atrás: L.T x = y
    x = np.zeros(n, dtype=float)
    for i in range(n - 1, -1, -1):
        x[i] = (y[i] - np.dot(L[i+1:, i], x[i+1:])) / L[i, i]

    return ________, ________

```


<!-- grade_id: q3_solution, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: q3_tests, points: 10, grade: true, solution: false -->
```python
A3 = np.array([[4., 2., 2.],
               [2., 5., 1.],
               [2., 1., 3.]])
b3 = np.array([8., 8., 6.])

L3, x3 = cholesky_solve(A3, b3)

assert np.allclose(np.triu(L3, 1), 0)
assert x3.shape == (3,)


```

## Pregunta 4 — QR mediante Gram-Schmidt modificado (10 puntos)

Implemente `mgs_qr(A)` usando **Gram-Schmidt modificado** para una matriz $A\in\mathbb{R}^{m\times n}$, con $m\ge n$ y columnas linealmente independientes.

Debe regresar matrices $Q$ y $R$ tales que

$
A=QR,
$

con

$
Q^TQ=I
$

y $R$ triangular superior.

No use `np.linalg.qr`.

## Código Guía
```Python
def mgs_qr(A):
    A = np.asarray(______, dtype=float)
    m, n = A.shape
    Q = np.zeros((______, ______), dtype=float)
    R = np.zeros((______, ______), dtype=float)
    V = A.copy()

    for i in range(n):
        R[i, i] = np.linalg.norm(V[:, i])
        Q[:, i] = V[:, i] / R[i, i]

        for j in range(i + 1, n):
            R[i, j] = np.dot(Q[:, i], V[:, j])
            V[:, j] = V[:, j] - R[i, j] * Q[:, i]

    return ______, ______

```

<!-- grade_id: q4_solution, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: q4_tests, points: 10, grade: true, solution: false -->
```python
A4 = np.array([[1., 1.],
               [1., 0.],
               [0., 1.]])

Q4, R4 = mgs_qr(A4)

assert Q4.shape == (3, 2)
assert R4.shape == (2, 2)
assert np.allclose(np.tril(R4, -1), 0)


```

## Entrega

Antes de entregar:

1. Reinicie el kernel.
2. Ejecute todas las celdas.
3. Verifique que no existan errores en las pruebas visibles.
4. Guarde el notebook con su nombre y número de cuenta.
