# Métodos Matemáticos Computacionales para Ciencia de Datos
## Laboratorio 2. Factorizaciones matriciales y matriz pseudoinversa
### Facultad de Estudios Superiores Acatlán · Universidad Nacional Autónoma de México

**Curso:** CdD 0503  
**Profesor:** Dr. Javier Orduz  
**Semestre:** 2027-1

---

### Datos del estudiante

**Nombre:** _______________________________________________

**Número de cuenta:** ______________________________________

**Fecha:** _________________________________________________


## Flujo general del laboratorio

<p align="center">
  <img src="./images/workflow.svg" width="1050" alt="Flujo general del Laboratorio 2">
</p>


# Instrucciones generales

Este laboratorio contiene **10 ejercicios** y tiene un valor total de **68 puntos**.

1. No elimine, agregue ni reordene celdas. `nbgrader` utiliza la estructura y los identificadores del notebook.
2. Ejecute las celdas en el orden en que aparecen.
3. Las celdas con `assert` contienen pruebas visibles; otras pruebas permanecen ocultas.
4. En los ejercicios con **Código guía**, copie el bloque mostrado, péguelo en la celda de respuesta y complete únicamente los espacios indicados.
5. Modifique solo las celdas que contienen:

```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

6. Sustituya `raise NotImplementedError()` por su solución.
7. Guarde frecuentemente el notebook y ejecútelo completo antes de entregarlo.
8. Las preguntas conceptuales asociadas se responderán por separado en Moodle.


# Objetivos

Al finalizar el laboratorio podrá:

- implementar la factorización LU mediante el esquema de Crout;
- resolver sistemas lineales usando sustitución hacia adelante y hacia atrás;
- implementar la factorización de Cholesky para matrices simétricas definidas positivas;
- construir una factorización QR mediante Gram–Schmidt modificado;
- resolver sistemas y problemas de mínimos cuadrados utilizando QR;
- calcular e interpretar la matriz pseudoinversa.


```python
import numpy as np

np.set_printoptions(precision=6, suppress=True)

```

## Funciones auxiliares

Las siguientes funciones implementan sustitución hacia adelante y hacia atrás. Se proporcionan como herramientas para concentrar el trabajo en las factorizaciones.


```python
def sustitucion_adelante(L, b):
    """Resuelve Ly=b para L triangular inferior."""
    L = np.asarray(L, dtype=float)
    b = np.asarray(b, dtype=float)
    n = len(b)
    y = np.zeros(n, dtype=float)

    for i in range(n):
        y[i] = (b[i] - L[i, :i] @ y[:i]) / L[i, i]

    return y


def sustitucion_atras(U, y):
    """Resuelve Ux=y para U triangular superior."""
    U = np.asarray(U, dtype=float)
    y = np.asarray(y, dtype=float)
    n = len(y)
    x = np.zeros(n, dtype=float)

    for i in range(n - 1, -1, -1):
        x[i] = (y[i] - U[i, i + 1:] @ x[i + 1:]) / U[i, i]

    return x

```

# Ejercicio 1. Preparación de matrices y propiedades básicas — 4 puntos

Considere las matrices:

$$
A_C=
\begin{bmatrix}
4&2&0\\
4&4&2\\
2&2&3
\end{bmatrix},
\qquad
A_S=
\begin{bmatrix}
4&1&1\\
1&3&0\\
1&0&2
\end{bmatrix}.
$$

Determine mediante NumPy si `A_C` es cuadrada, si `A_S` es simétrica y si `A_S` es definida positiva.


```python
A_C = np.array([
    [4., 2., 0.],
    [4., 4., 2.],
    [2., 2., 3.]
])

A_S = np.array([
    [4., 1., 1.],
    [1., 3., 0.],
    [1., 0., 2.]
])

```

## Código guía

```python
es_cuadrada_C = ...
es_simetrica_S = ...
valores_propios_S = ...
es_definida_positiva_S = ...
```


<!-- grade_id: cell-e86ad5a50d3046e7, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-6d15736c18ab4501, points: 4, grade: true, solution: false -->
```python
assert es_cuadrada_C
assert es_simetrica_S


```

# Ejercicio 2. Factorización LU mediante Crout — 10 puntos

En la versión de Crout utilizada aquí:

$$
A=LU,
$$

donde \(L\) es triangular inferior y la diagonal de \(U\) está formada por 1s, "unos".

<p align="center">
  <img src="./images/crout_flow.svg" width="950" alt="Flujo de la factorización de Crout">
</p>


## Código guía

```python
def crout(A):
    A = np.asarray(A, dtype=float)
    n = A.shape[0]
    L = np.zeros_like(A)
    U = np.eye(n)

    for j in range(n):
        for i in range(j, n):
            L[i, j] = ______________________________

        if abs(L[j, j]) < 1e-14:
            raise np.linalg.LinAlgError("Pivote nulo en Crout")

        for i in range(j + 1, n):
            U[j, i] = ______________________________

    return L, U
```


<!-- grade_id: cell-7bcb51ba547042d0, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-a9ae9af0921e4f46, points: 10, grade: true, solution: false -->
```python
L_C, U_C = crout(A_C)

assert np.allclose(np.diag(U_C), 1.0)
assert np.allclose(L_C @ U_C, A_C)


```

# Ejercicio 3. Solución de un sistema con Crout — 6 puntos

Use $A_C=LU$ para resolver:

$$
A_Cx=b_C,
\qquad
b_C=
\begin{bmatrix}
2\\4\\6
\end{bmatrix}.
$$

Primero resuelva $Ly=b_C$ y después $Ux=y$.


```python
b_C = np.array([2., 4., 6.])
```

## Código guía

```python
L_C, U_C = crout(A_C)
y_C = ...
x_C = ...
residuo_C = ...
```


<!-- grade_id: cell-e2719e10c28e4e32, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-619bcaf0635d400b, points: 6, grade: true, solution: false -->
```python
assert y_C.shape == (3,)
assert x_C.shape == (3,)

```

# Ejercicio 4. Factorización de Cholesky — 8 puntos

Para una matriz real, simétrica y definida positiva:

$$
A=LL^{T},
$$

donde $L$ es triangular inferior con diagonal positiva.

<p align="center">
  <img src="./images/cholesky_flow.svg" width="950" alt="Flujo de la factorización de Cholesky">
</p>


## Código guía

```python
def cholesky_manual(A):
    ...
    for i in range(n):
        for j in range(i + 1):
            suma = ______________________

            if i == j:
                L[i, j] = ______________________
            else:
                L[i, j] = ______________________

    return L
```


<!-- grade_id: cell-b819177b45a542c2, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-86043c02edde4b94, points: 8, grade: true, solution: false -->
```python
L_S = cholesky_manual(A_S)

assert np.allclose(L_S, np.tril(L_S))
assert np.all(np.diag(L_S) > 0)


```

# Ejercicio 5. Solución de un sistema con Cholesky — 6 puntos

Resuelva:

$$
A_Sx=b_S,
\qquad
b_S=
\begin{bmatrix}
7\\8\\3
\end{bmatrix}.
$$

Use $A_S=LL^T,$ sustitución hacia adelante y sustitución hacia atrás.


```python
b_S = np.array([7., 8., 3.])
```

## Código guía

```python
L_S = cholesky_manual(A_S)
y_S = ...
x_S = ...
residuo_S = ...
```


<!-- grade_id: cell-4db53d5f06d749e2, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-094c206051294dcf, points: 6, grade: true, solution: false -->
```python
assert x_S.shape == (3,)
assert residuo_S >= 0


```

# Ejercicio 6. QR mediante Gram–Schmidt modificado — 10 puntos

Para una matriz $A\in\mathbb{R}^{m\times n},$ con $m\ge n$ y columnas linealmente independientes, buscamos:

$$
A=QR,
$$

donde $Q^TQ=I$ y $R$ es triangular superior.

<p align="center">
  <img src="./images/gram_schmidt_qr.svg" width="1000" alt="Gram-Schmidt modificado y factorización QR">
</p>


```python
A_QR = np.array([
    [1., 1., 0.],
    [1., 0., 1.],
    [0., 1., 1.],
    [1., 1., 1.]
])
```

## Código guía

El siguiente código debe ser completado con otras funciones, condiciones u operadores.

```python
def gram_schmidt_modificado(A):
    A = np.asarray(A, dtype=float)
    m, n = A.shape
    Q = np.zeros((m, n))
    R = np.zeros((n, n))
    V = A.copy()

    for i in range(n):
        R[i, i] = ______________________
        Q[:, i] = ______________________

        for j in range(i + 1, n):
            R[i, j] = ______________________
            V[:, j] = ______________________

    return Q, R
```


<!-- grade_id: cell-a9c611aa83cf45ed, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-87575e5b42a64ad4, points: 10, grade: true, solution: false -->
```python
Q_QR, R_QR = gram_schmidt_modificado(A_QR)

assert Q_QR.shape == (4, 3)
assert R_QR.shape == (3, 3)


```

# Ejercicio 7. Solución de mínimos cuadrados mediante QR — 6 puntos

Para:

$$
\min_x \|A_{QR}x-b_{QR}\|_2,
$$

use $A_{QR}=QR$ y resuelva:

$$
Rx=Q^Tb_{QR}.
$$


```python
b_QR = np.array([2., 1., 2., 2.5])
```

## Código guía

```python
Q_QR, R_QR = gram_schmidt_modificado(A_QR)
c_QR = ...
x_QR = ...
residuo_QR = ...
```


<!-- grade_id: cell-07ccbfc20b6947bc, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-ff3960b4f8454b0f, points: 6, grade: true, solution: false -->
```python
assert c_QR.shape == (3,)
assert x_QR.shape == (3,)


```

# Ejercicio 8. Pseudoinversa mediante QR — 8 puntos

Si $A\in\mathbb{R}^{m\times n}$ tiene rango columna completo y $A=QR,$ entonces:

$$
A^{+}=R^{-1}Q^T.
$$

Implemente esta fórmula sin usar `np.linalg.pinv` dentro de la función.

<p align="center">
  <img src="./images/pseudoinverse_flow.svg" width="1000" alt="Construcción de la pseudoinversa">
</p>


## Código guía

```python
def pseudoinversa_qr(A):
    Q, R = gram_schmidt_modificado(A)
    R_inv = ______________________
    return ______________________
```


<!-- grade_id: cell-2b12fae1eb0a4ae4, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-9f6ec035c36e44da, points: 8, grade: true, solution: false -->
```python
A_QR_pinv = pseudoinversa_qr(A_QR)

assert A_QR_pinv.shape == (3, 4)
assert np.allclose(A_QR_pinv @ A_QR, np.eye(3), atol=1e-10)


```

# Ejercicio 9. Ajuste lineal con la pseudoinversa — 6 puntos

Ajuste el modelo:

$$
y\approx \beta_0+\beta_1t
$$

a los datos dados. Construya la matriz de diseño y obtenga:

$$
\beta=X^{+}y.
$$


```python
t_datos = np.array([0., 1., 2., 3., 4.])
y_datos = np.array([1.1, 2.0, 2.9, 4.2, 4.8])
```

## Código guía

```python
X_diseno = ...
beta = ...
y_estimado = ...
error_ajuste = ...
```


<!-- grade_id: cell-0202779ed576433f, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-15fe88f692aa496b, points: 6, grade: true, solution: false -->
```python
assert X_diseno.shape == (5, 2)
assert beta.shape == (2,)


```

# Ejercicio 10. Pseudoinversa de una matriz con rango deficiente — 4 puntos

La fórmula QR del ejercicio anterior requiere rango columna completo. Para una matriz con columnas dependientes usaremos `np.linalg.pinv`, basada en la descomposición en valores singulares.

Considere:

$$
A_R=
\begin{bmatrix}
1&2\\
2&4\\
3&6
\end{bmatrix}.
$$

Calcule su rango, su pseudoinversa y verifique numéricamente dos identidades de Moore–Penrose.


```python
A_R = np.array([
    [1., 2.],
    [2., 4.],
    [3., 6.]
])
```

## Código guía

```python
rango_R = ...
A_R_pinv = ...
identidad_1 = ...
identidad_2 = ...
```


<!-- grade_id: cell-e4d73a96260b45d2, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-e463a545c4284eb2, points: 4, grade: true, solution: false -->
```python
assert rango_R == 1
assert A_R_pinv.shape == (2, 3)


```

# Cierre

Ha completado el Laboratorio 2. En este notebook implementó y utilizó:

- Crout;
- Cholesky;
- Gram–Schmidt modificado;
- factorización QR;
- solución de mínimos cuadrados;
- matriz pseudoinversa.

Las preguntas conceptuales y de interpretación se encuentran en Moodle.
