# Métodos Matemáticos Computacionales para Ciencia de Datos
## Laboratorio 1. Solución Numérica de Sistemas Lineales
### Facultad de Estudios Superiores Acatlán · Universidad Nacional Autónoma de México

**Curso:** CdD 0503  
**Profesor:** Dr. Javier Orduz  
**Semestre:** 2027-1

---

### Nombre del estudiante

**Nombre:** __________________________________________

**Número de cuenta:** _________________________________

**Fecha:** ____________________________________________

---

## Flujo general del laboratorio


# Instrucciones generales

Bienvenido(a) al laboratorio.

Este laboratorio contiene **10 Ejercicios** que deberán resolverse de manera individual utilizando **Python**, **NumPy** y **Jupyter Notebook**.

Antes de comenzar, lea cuidadosamente las siguientes instrucciones.

## Reglas generales

1. **No elimine, agregue o reordene celdas del notebook.**
   La estructura del laboratorio es utilizada por **nbgrader** para realizar la evaluación automática.

2. **No modifique las celdas marcadas como "Read-only".**
   Estas contienen información, código base o pruebas de evaluación.

3. **Las celdas con `assert` forman parte de la evaluación automática.**
   Algunas pruebas son visibles y otras permanecen ocultas. Si un `assert` falla, revise cuidadosamente su solución antes de continuar.

4. **Resuelva únicamente las celdas que contienen el siguiente código:**

```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

Únicamente debe modificar el código a partir de esa línea.

5. Algunos ejercicios incluyen un **Código guía**.

En esos casos:

- copie el código guía,
- péguelo en la celda de respuesta,
- complete únicamente las partes faltantes.

No es necesario modificar el algoritmo completo.

6. Ejecute las celdas en el orden en que aparecen.

7. Guarde frecuentemente su notebook.

8. Antes de entregar el laboratorio, asegúrese de que **todas las celdas hayan sido ejecutadas correctamente**.

---

<center>
<img src="../../images/lab01/workflow.svg" width="1000">
</center>


# Objetivos del laboratorio

Al finalizar este laboratorio usted será capaz de:

- Representar sistemas lineales utilizando matrices y vectores.
- Resolver sistemas mediante métodos directos.
- Implementar el método de Jacobi.
- Implementar el método de Gauss-Seidel.
- Analizar la convergencia de métodos iterativos.
- Comparar métodos directos e iterativos.
- Interpretar resultados numéricos mediante tablas y gráficas.

---

## Entregables

Este laboratorio consta de **10 Ejercicios**.

Cada Ejercicio puede contener:

- preguntas conceptuales;
- programación en Python;
- análisis de resultados;
- interpretación de gráficas.

---

¡Mucho éxito!

## Solución numérica de sistemas lineales

### Métodos de Jacobi y Gauss-Seidel

Este laboratorio introduce la solución numérica de sistemas lineales mediante métodos iterativos.

Trabajaremos con sistemas de la forma:

$$
Ax=b
$$

### Diagramas de los algoritmos
<center>
<img src="../../images/lab01/jacobi.svg" width="250">
<img src="../../images/lab01/gauss_seidel.svg" width="850">
</center>

### ¿Por qué converge Jacobi?

La convergencia de un método iterativo depende de las propiedades de la matriz del sistema lineal.

Una condición suficiente ampliamente utilizada es que la matriz sea **estrictamente diagonalmente dominante**. En ese caso, tanto el método de Jacobi como el de Gauss-Seidel convergen para cualquier aproximación inicial.

El siguiente diagrama resume esta idea.

<center>
<img src="../../images/lab01/convergence.svg" width="900">
</center>



## Resumen conceptual

<p align="center">
  <img src="../../images/lab01/summary.svg" width="1000">
</p>

```python
import numpy as np
import matplotlib.pyplot as plt
```

# Ejercicio 1. Representación matricial

Considere el sistema lineal:

$$
A =
\begin{bmatrix}
4 & -1 & 0 \\
-1 & 4 & -1 \\
0 & -1 & 3
\end{bmatrix},
\qquad
b =
\begin{bmatrix}
15 \\
10 \\
10
\end{bmatrix}.
$$

Complete las variables solicitadas para identificar propiedades básicas del sistema.

```python
A = np.array([
    [4, -1, 0],
    [-1, 4, -1],
    [0, -1, 3]
], dtype=float)

b = np.array([15, 10, 10], dtype=float)
```

## Actividad 1

Complete el siguiente código para obtener:

- Las dimensiones de la matriz `A`.
- El número de ecuaciones.
- El número de incógnitas.
- Si la matriz es cuadrada.
- La norma Euclidiana del vector `b`.

**No modifique el nombre de las variables.**

<!-- grade_id: cell-11d3c3991f1e2be7, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-9a17c18038529177, points: 5, grade: true, solution: false -->
```python
# Pruebas básicas (visibles)

assert shape_A == (3, 3)
assert n == 3

```

# Ejercicio 2. Solución exacta del sistema

Utilizaremos `numpy.linalg.solve()` para obtener la solución exacta del sistema.

Esta solución servirá como referencia para comparar los métodos iterativos.

<!-- grade_id: cell-09265ba19e3068e1, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-32c318641a967433, points: 5, grade: true, solution: false -->
```python
# Pruebas básicas (visibles)

assert np.allclose(A @ x_exacta, b)
assert residuo < 1e-12

```

# Ejercicio 3. Primera iteración del método de Jacobi

El método de Jacobi construye una nueva aproximación utilizando únicamente los valores de la iteración anterior.

En este ejercicio realizará **una sola iteración** del método utilizando como aproximación inicial

$$
x^{(0)}=
\begin{bmatrix}
0\\
0\\
0
\end{bmatrix}.
$$

Calcule manualmente las tres componentes del vector

$$
\vec{x}^{(1)}=(x^{(1)}_1,x^{(1)}_2,x^{(1)}_3).
$$

No utilice ciclos (`for`) en este ejercicio.

```python
# Aproximación inicial

x0 = np.zeros(3)

print("x^(0) =")
print(x0)
```

## Ecuaciones del método de Jacobi

Para aplicar el método de Jacobi, primero despejamos cada variable del sistema de ecuaciones.

Dado el sistema

$$
\begin{aligned}
4x_1-x_2 &=15\\
-x_1+4x_2-x_3 &=10\\
-x_2+3x_3 &=10
\end{aligned}
$$

obtenemos las siguientes ecuaciones iterativas:

$$
\begin{align}
x_1^{(k+1)}&=\frac{15+x_2^{(k)}}{4},\\
x_2^{(k+1)}&=\frac{10+x_1^{(k)}+x_3^{(k)}}{4},\\
x_3^{(k+1)}&=\frac{10+x_2^{(k)}}{3}.
\end{align}
$$

Observe que **todas las componentes de $x^{(k+1)}$** se calculan utilizando únicamente los valores de la iteración anterior $x^{(k)}.$ **Implemente la primera iteración de Jacobi** en la siguiente celda.

### Pregunta

¿Qué ocurriría si, durante el cálculo de $x_2^{(1)}$, utilizáramos el nuevo valor de $x_1^{(1)}$ en lugar del valor anterior $x_1^{(0)}$?

No responda todavía. Analizaremos esa diferencia cuando estudiemos el método de Gauss-Seidel.

<!-- grade_id: cell-974d2ca3dfa9e38f, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-ce05c2b788fa3c16, points: 4, grade: true, solution: false -->
```python
x1_correcta = np.array([
    15/4,
    10/4,
    10/3
])

# Prueba básica (visible)

assert x1.shape == (3,)

```

# Ejercicio 4. Implementación del método de Jacobi

En el Ejercicio anterior calculó una sola iteración del método de Jacobi.

Ahora automatizará el proceso para realizar varias iteraciones hasta que la solución converja.

Utilizaremos como criterio de convergencia

$$
\|x^{(k+1)}-x^{(k)}\|<10^{-6}.
$$

Además, almacenaremos el error de cada iteración para analizar posteriormente la convergencia del método.

```python
tol = 1e-6
max_iter = 100

x0 = np.zeros(3)
```

## Actividad

Complete la función `jacobi`.

La función debe regresar dos objetos:

- La aproximación final.
- Una lista con el error de cada iteración.

```python
def jacobi(A, b, x0, tol=1e-6, max_iter=100):

    n = len(b)

    x = x0.copy()

    errores = []

    for k in range(__________):

        x_new = np.zeros_like(__________)

        for i in range(__________):

            suma = np.sum(A[i, :] * x) - A[i, i] * x[i]

            x_new[i] = (b[i] - suma) / A[i, i]

        error = np.linalg.norm(__________-__________)

        errores.append(__________)

        if error < tol:

            return x_new, errores

        x = x_new.copy()

    return __________, __________
```
Escriba su solución de la función `jacobi` en la siguiente celda.

<!-- grade_id: cell-8e8358cb829f56cd, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

```python
x_jacobi, errores = jacobi(A, b, x0)

print(x_jacobi)

print(f"Iteraciones: {len(errores)}")

print(f"Error final: {errores[-1]:.2e}")
```

<!-- grade_id: cell-ba5f3f08dce8a822, points: 8, grade: true, solution: false -->
```python
# Pruebas básicas (visibles)

assert isinstance(errores, list)
assert len(errores) > 0

```

# Ejercicio 5. Gráfica de convergencia de Jacobi

En este ejercicio debes **graficar** el error del método de Jacobi en función del número de iteración.

La lista `errores` fue generada en el ejercicio anterior. Cada elemento representa:

$$
\|x^{(k+1)}-x^{(k)}\|
$$

Usaremos una escala logarítmica en el eje vertical para observar mejor cómo disminuye el error.

## Código guía

Complete el código para producir una gráfica usando `plt.semilogy`.

```python
iteraciones_jacobi = np.arange(1, len(__________) + 1)

fig, ax = plt.subplots(figsize=(6, 4))

ax.semilogy(__________,__________, marker="o")

ax.set_xlabel("Iteración")
ax.set_ylabel("Error")
ax.set_title("Convergencia del método de Jacobi")
ax.grid(True)

plt.show()
```

<!-- grade_id: cell-c4852be8279102d4, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-efdfe72ad0f74487, points: 3, grade: true, solution: false -->
```python
# Pruebas básicas (visibles)

assert len(iteraciones_jacobi) == len(errores)
assert iteraciones_jacobi[0] == 1

```

# Ejercicio 6. Dominancia diagonal

No todas las matrices garantizan la convergencia del método de Jacobi.

Una condición suficiente de convergencia es que la matriz sea **diagonalmente dominante por filas**.

Una matriz

$$
A=[a_{ij}]
$$

es diagonalmente dominante si para cada fila se cumple

$$
\|a_{ii}\|
\ge
\sum_{j\neq i}\|a_{ij}\|.
$$

En este ejercicio implementará una función que verifique esta propiedad.

## Código guía

Complete las líneas faltantes del siguiente algoritmo.

```python
def es_diagonalmente_dominante(A):

    n = len(A)

    for i in range(n):

        diagonal = ______________________

        suma = __________________________

        if diagonal < suma:

            return False

    return True
```

<!-- grade_id: cell-89b0f895b8961825, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

```python
print("¿A es diagonalmente dominante?")

print(es_diagonalmente_dominante(A))
```

<!-- grade_id: cell-a98d21b1a39814e4, points: 4, grade: true, solution: false -->
```python
# Caso del laboratorio
assert es_diagonalmente_dominante(A)

```

# Ejercicio 7. Implementación del método de Gauss-Seidel

El método de Gauss-Seidel es similar al método de Jacobi, pero con una diferencia importante:

- Jacobi calcula todas las nuevas componentes usando únicamente $x^{(k)}.$
- Gauss-Seidel usa inmediatamente los valores nuevos conforme se van calculando.

Esto significa que, al calcular $x_i^{(k+1)},$ ya se pueden usar los valores actualizados de $x_1^{(k+1)}, x_2^{(k+1)}, \ldots, x_{i-1}^{(k+1)}.$

## Código guía

Complete las líneas faltantes del siguiente algoritmo.

```python
def gauss_seidel(A, b, x0, tol=1e-6, max_iter=100):

    n = len(b)
    x = x0.copy()
    errores = []

    for k in range(max_iter):

        x_old = x.copy()

        for i in range(n):

            suma = ______________________

            x[i] = ______________________

        error = ______________________

        errores.append(error)

        if error < tol:
            return x, errores

    return x, errores
```

<!-- grade_id: cell-5e5bb2b2d58545a8, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

```python
x_gs, errores_gs = gauss_seidel(A, b, x0)

print("Solución aproximada con Gauss-Seidel:")
print(x_gs)

print(f"Iteraciones: {len(errores_gs)}")
print(f"Error final: {errores_gs[-1]:.2e}")
```

<!-- grade_id: cell-285ce5b39bfef6e1, points: 6, grade: true, solution: false -->
```python
# Pruebas básicas (visibles)

assert isinstance(errores_gs, list)
assert len(errores_gs) > 0

```

# Ejercicio 8. Comparación entre Jacobi y Gauss-Seidel

En este ejercicio comparará el comportamiento de los métodos de Jacobi y Gauss-Seidel.

Analizará:

- Número de iteraciones.
- Error final.
- Curva de convergencia.
- Tiempo de ejecución (opcional).

El objetivo es identificar las principales diferencias entre ambos métodos.

## Código guía

Complete el siguiente código para construir una tabla comparativa.

```python
resultados = {
    "Método": ["Jacobi", "Gauss-Seidel"],
    "Iteraciones": [__________, __________],
    "Error final": [__________, __________]
}
```

<!-- grade_id: cell-042af658defae839, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

```python
import pandas as pd

df = pd.DataFrame(resultados)

df
```

## Código guía

Complete la gráfica para comparar la convergencia de ambos métodos.

```python
plt.figure(figsize=(7,4))

plt.semilogy(______________,______________,______________)

plt.semilogy(______________,______________,______________)

plt.xlabel("Iteración")
plt.ylabel("Error")
plt.title("Comparación de convergencia")

plt.legend()
plt.grid(True)
plt.show()
```

<!-- grade_id: cell-fe2bf8e00aba2920, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-5325aedae4f33a4f, points: 10, grade: true, solution: false -->
```python
# Pruebas básicas visibles

assert isinstance(resultados, dict)
assert "Método" in resultados
assert "Iteraciones" in resultados
assert "Error final" in resultados

assert len(resultados["Método"]) == 2
assert len(resultados["Iteraciones"]) == 2
assert len(resultados["Error final"]) == 2

```

# Ejercicio 9. Un caso donde Jacobi no converge

Hasta ahora trabajamos con una matriz para la cual los métodos iterativos convergen.

En este ejercicio analizaremos un sistema diferente donde el método de Jacobi puede fallar.

El objetivo es observar que los métodos iterativos no siempre convergen y que las propiedades de la matriz son importantes.

```python
A_nc = np.array([
    [1, 3],
    [2, 1]
], dtype=float)

b_nc = np.array([5, 5], dtype=float)

x0_nc = np.zeros(2)
```

## Actividad 1

Verifique si la matriz `A_nc` es diagonalmente dominante usando la función que ya implementó.

<!-- grade_id: cell-e80d2d74417b8aa3, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: cell-29e9de3d045947ab, points: 5, grade: true, solution: false -->
```python
assert dominante_nc is False
```

## Actividad 2

Aplique el método de Jacobi a este nuevo sistema utilizando exactamente los mismos parámetros empleados anteriormente.

Observe el comportamiento del error durante las iteraciones.

<!-- grade_id: cell-2ba78b103955acbf, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

```python
print("Número de iteraciones:", len(errores_nc))
print("Último error:", errores_nc[-1])
print("Aproximación final:")
print(x_nc)
```

<!-- grade_id: cell-85ed58ed08e4ab12, points: 5, grade: true, solution: false -->
```python
assert isinstance(errores_nc, list)
assert len(errores_nc) == 20
```

## Actividad 3

Grafique el error obtenido durante las iteraciones.

```python
plt.figure(figsize=(6,4))

plt.semilogy(errores_nc, marker="o")

plt.xlabel("Iteración")
plt.ylabel("Error")
plt.title("Jacobi: caso sin convergencia")

plt.grid(True)

plt.show()
```

# Ejercicio 10. Conclusión del laboratorio

En este último ejercicio comparará las soluciones obtenidas con Jacobi y Gauss-Seidel contra la solución exacta.

El objetivo es cerrar el laboratorio verificando que ambos métodos produjeron aproximaciones correctas para el sistema original.

<!-- grade_id: cell-ce2bce3276be7f71, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

```python
print(f"Error Jacobi vs solución exacta: {error_jacobi_exacta:.2e}")
print(f"Error Gauss-Seidel vs solución exacta: {error_gs_exacta:.2e}")
print(f"Método con menos iteraciones: {mejor_metodo}")
```

<!-- grade_id: cell-a54757b9e3292c07, points: 3, grade: true, solution: false -->
```python
# Pruebas básicas (visibles)

assert mejor_metodo in ["Jacobi", "Gauss-Seidel"]

```

## Reflexión final

Responda de manera breve y clara:

1. ¿Cuál es la diferencia fundamental entre Jacobi y Gauss-Seidel?
2. ¿Por qué la dominancia diagonal favorece la convergencia?
3. Mencione una aplicación real donde aparezcan sistemas lineales grandes.


<!-- grade_id: cell-reflexion-final-lab01, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

# ¡Felicidades!

Ha completado el Laboratorio 1.

En este laboratorio aprendió a:

- Representar sistemas lineales con NumPy.
- Resolver sistemas mediante un método directo.
- Implementar el método de Jacobi.
- Analizar su convergencia.
- Implementar el método de Gauss-Seidel.
- Comparar ambos métodos.
- Identificar casos donde Jacobi no converge.

En el siguiente laboratorio estudiaremos métodos de factorización matricial y su aplicación en la solución de sistemas lineales.