# Métodos Matemáticos Computacionales para Ciencia de Datos
## Laboratorio 3. Diagonalización, valores propios, SVD y pseudoinversa

**Valor total:** 68 puntos  
**Número de ejercicios:** 9

Nombre: ____________________________________  
Número de cuenta: ___________________________  
Fecha: _____________________________________


## Flujo general del laboratorio

<p align="center">
  <img src="./images/workflow.svg" width="1050">
</p>

En este laboratorio se conectan cuatro ideas:

$$
A\mathbf v=\lambda\mathbf v
\;\longrightarrow\;
A=PDP^{-1}
\;\longrightarrow\;
A=U\Sigma V^T
\;\longrightarrow\;
A^+=V\Sigma^+U^T.
$$

El objetivo es estudiar estas descomposiciones **numéricamente** y reconocer cuándo existen las formas requeridas.

# Instrucciones generales

Este laboratorio contiene **9 ejercicios** y tiene un valor total de **68 puntos**.

- No agregue ni elimine celdas.
- No modifique las celdas de solo lectura.
- Modifique únicamente las celdas que contienen `# ESCRIBA SU CÓDIGO AQUÍ`.
- Ejecute las celdas en orden.
- Los `assert` forman parte de la evaluación automática.
- Los tests visibles comprueban propiedades básicas; la evaluación incluye comprobaciones adicionales.
- No escriba valores esperados de forma manual: implemente el procedimiento solicitado.
- Antes de entregar, reinicie el kernel y ejecute **Run All**.
- Guarde el notebook con frecuencia.


# Objetivos

Al finalizar el laboratorio podrá:

- obtener numéricamente valores y vectores propios;
- verificar si una matriz admite una diagonalización completa;
- construir y validar $A=PDP^{-1}$;
- utilizar la diagonalización para calcular potencias matriciales;
- calcular e interpretar una descomposición en valores singulares;
- estudiar rango y aproximaciones mediante valores singulares;
- construir la pseudoinversa de Moore–Penrose mediante SVD;
- utilizar la pseudoinversa en sistemas compatibles, incompatibles y de rango deficiente.


```python
import numpy as np

np.set_printoptions(precision=6, suppress=True)
TOL = 1e-10

```

# Ejercicio 1. Obtención numérica de valores y vectores propios — 8 puntos

Considere

$$
A_1=
\begin{pmatrix}
2&1\\
1&2
\end{pmatrix}.
$$

Use `np.linalg.eigh`, apropiado para matrices reales simétricas, para obtener sus valores y vectores propios.

Guarde:
- los valores propios en `eigvals_1`;
- los vectores propios, por columnas, en `eigvecs_1`;
- el residuo máximo $\max_i\|A_1v_i-\lambda_i v_i\|_2$ en `eig_residual_1`.


```python
A_1 = np.array([
    [2., 1.],
    [1., 2.]
])

```

<!-- grade_id: answer-1e3dd5104b29, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-a10740928b30, points: 8, grade: true, solution: false -->
```python
assert eigvals_1.shape == (2,)
assert eigvecs_1.shape == (2, 2)
assert np.all(np.diff(eigvals_1) >= 0)


```

# Ejercicio 2. ¿Existe una diagonalización? — 6 puntos

Implemente `es_diagonalizable(A, tol=1e-10)`.

La función debe:

1. comprobar que \(A\) sea cuadrada;
2. calcular sus valores y vectores propios con `np.linalg.eig`;
3. determinar el rango de la matriz de vectores propios;
4. regresar `True` únicamente cuando existan \(n\) vectores propios linealmente independientes.

> **Nota:** en este ejercicio se utilizan matrices reales con espectro real.


<!-- grade_id: answer-b9cac5ce354e, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-1b097a3292f2, points: 6, grade: true, solution: false -->
```python
assert es_diagonalizable(A_1)


```

# Ejercicio 3. Construcción de \(P\), \(D\) y verificación — 8 puntos

Para $A_1$, construya

$$
A_1=PDP^{-1}.
$$

Use los resultados del Ejercicio 1:

$$
P=[v_1\;v_2],
\qquad
D=\operatorname{diag}(\lambda_1,\lambda_2).
$$

Calcule también

$$
\|A_1-PDP^{-1}\|_2.
$$

Guarde el resultado en `diag_residual_1`.


<!-- grade_id: answer-0eb71e0cda73, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-b26dc7393458, points: 8, grade: true, solution: false -->
```python
assert P_1.shape == A_1.shape
assert D_1.shape == A_1.shape
assert np.allclose(D_1, np.diag(np.diag(D_1)))


```

# Ejercicio 4. Potencias mediante diagonalización — 6 puntos

Utilice

$$
A^k=PD^kP^{-1}
$$

para calcular $A_1^{10}$ **sin utilizar `np.linalg.matrix_power(A_1, 10)` para construir el resultado**.

Guarde la matriz en `A1_power10_diag`.

Puede utilizar `np.linalg.matrix_power` únicamente para calcular `A1_power10_reference` y verificar el resultado.


<!-- grade_id: answer-247326b45909, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-a6a8e72e2c4f, points: 6, grade: true, solution: false -->
```python
assert A1_power10_diag.shape == (2, 2)
assert power_error_1 >= 0


```

# Ejercicio 5. Matriz no diagonalizable — 6 puntos

Considere la matriz de Jordan

$$
A_J=
\begin{pmatrix}
2&1\\
0&2
\end{pmatrix}.
$$

Obtenga numéricamente sus valores y vectores propios y determine:

- `rank_eigvecs_J`: rango de la matriz de vectores propios;
- `diagonalizable_J`: resultado usando su función `es_diagonalizable`;
- `num_independent_J`: número de vectores propios linealmente independientes detectados.

El objetivo es distinguir entre **tener eigenvalores** y **tener suficientes eigenvectores para formar una base**.


```python
A_J = np.array([
    [2., 1.],
    [0., 2.]
])

```

<!-- grade_id: answer-0b1892582e78, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-1a1402e62fb4, points: 6, grade: true, solution: false -->
```python
assert eigvals_J.shape == (2,)
assert eigvecs_J.shape == (2, 2)


```

# Ejercicio 6. Descomposición en valores singulares — 10 puntos

Considere

$$
A_2=
\begin{pmatrix}
3&0\\
0&2\\
0&0
\end{pmatrix}.
$$

Calcule la SVD reducida:

$$
A_2=U\Sigma V^T.
$$

Use:

```python
U_2, s_2, Vt_2 = np.linalg.svd(A_2, full_matrices=False)
```

Construya explícitamente `Sigma_2` a partir de `s_2` y calcule:

$$
\|A_2-U\Sigma V^T\|_2.
$$


```python
A_2 = np.array([
    [3., 0.],
    [0., 2.],
    [0., 0.]
])

```

<!-- grade_id: answer-50ebff4caf20, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-139cd649f99a, points: 10, grade: true, solution: false -->
```python
assert U_2.shape == (3, 2)
assert s_2.shape == (2,)
assert Vt_2.shape == (2, 2)
assert Sigma_2.shape == (2, 2)


```

# Ejercicio 7. Rango numérico y aproximación SVD — 8 puntos

Considere

$$
A_3=
\begin{pmatrix}
3&0&0\\
0&1&0\\
0&0&0.01
\end{pmatrix}.
$$

1. Calcule su SVD.
2. Determine el rango numérico usando el umbral `tol_rank = 0.05`.
3. Construya la mejor aproximación de rango 2,

$$
A_{3,2}=U_{:,1:2}\Sigma_{1:2,1:2}V_{1:2,:}^T,
$$

adaptando correctamente los índices de Python.
4. Calcule el error espectral $\|A_3-A_{3,2}\|_2$.


```python
A_3 = np.diag([3., 1., 0.01])
tol_rank = 0.05
```

<!-- grade_id: answer-fcc4b897d0b5, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-62e62a8c2920, points: 8, grade: true, solution: false -->
```python
assert s_3.shape == (3,)
assert A3_rank2.shape == A_3.shape
assert rank_numeric_3 >= 0


```

# Ejercicio 8. Pseudoinversa mediante SVD — 10 puntos

Implemente

```python
pseudoinversa_svd(A, tol=1e-12)
```

utilizando

$$
A=U\Sigma V^T,
\qquad
A^+=V\Sigma^+U^T.
$$

Para cada valor singular:

$$
\sigma_i^+=
\begin{cases}
1/\sigma_i,& \sigma_i>\text{tol},\\
0,&\sigma_i\leq\text{tol}.
\end{cases}
$$

**No utilice `np.linalg.pinv` dentro de su función.**


<!-- grade_id: answer-5d98d2b07532, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-cbfc8f2e185d, points: 10, grade: true, solution: false -->
```python
assert A2_pinv_svd.shape == (2, 3)
assert np.all(np.isfinite(A2_pinv_svd))


```

# Ejercicio 9. Pseudoinversa y solución de un sistema de rango deficiente — 6 puntos

Considere

$$
A_4=
\begin{pmatrix}
1&1\\
2&2\\
3&3
\end{pmatrix},
\qquad
b_4=
\begin{pmatrix}
1\\2\\2
\end{pmatrix}.
$$

La matriz $A_4$ tiene rango deficiente.

Calcule:

$$
x^+=A_4^+b_4
$$

usando **su función `pseudoinversa_svd`**.

Después calcule:

$$
r=b_4-A_4x^+,
\qquad
\|r\|_2.
$$

Finalmente compruebe numéricamente la condición de ortogonalidad de mínimos cuadrados:

$$
A_4^Tr\approx0.
$$

Guarde su norma en `normal_residual_4`.


```python
A_4 = np.array([
    [1., 1.],
    [2., 2.],
    [3., 3.]
])

b_4 = np.array([1., 2., 2.])

```

<!-- grade_id: answer-8e2e2c2953d2, points: 0, grade: false, solution: true -->
```python
# ESCRIBA SU CÓDIGO AQUÍ
raise NotImplementedError()
```

<!-- grade_id: test-249590aeea78, points: 6, grade: true, solution: false -->
```python
assert A4_pinv.shape == (2, 3)
assert x4_plus.shape == (2,)
assert residual_4.shape == (3,)


```

# Resumen

En este laboratorio trabajó numéricamente con:

$$
A\mathbf v=\lambda\mathbf v,
\qquad
A=PDP^{-1},
\qquad
A=U\Sigma V^T,
\qquad
A^+=V\Sigma^+U^T.
$$

Antes de entregar:

1. reinicie el kernel;
2. ejecute todas las celdas;
3. verifique que no existan errores inesperados;
4. guarde el notebook.

**Puntuación total: 68 puntos.**


```python

```